# SLU 17_3 - Exam Prep I

## Batch 2 - Wave 2 Python Exam
This is the Python exam from batch 2 - wave 2. Exam Duration: 2h

```python
import json
from operator import itemgetter
import random
#used for evaluation
from utils import DataGenerator
import math
import sys

data = DataGenerator()
```

## Question 1

The variable `data` is a list of dictionaries.

Create a function `replace_unknown` that:
- Receives as argument a list of dictionaries;
- In each dictionary, it replaces all the '?' you find with the string value '1';
- Returns the modified list of dictionaires;

Note: values should retain the original datatypes.

<!-- grade_id: cell-7d5d16f632d819c3, points: 0, grade: false, solution: true -->
```python
def replace_unknown(data):
    replace = [{w: w2.replace('?', '1') if isinstance(w2, str) else w2 for w, w2 in animal.items()}
              for animal in data]
    return replace
```

<!-- grade_id: cell-2ec0be01019aa1ce, points: 3, grade: true, solution: false -->
```python
data_raw = data.data
columns = list(data_raw[0].keys())
data_replaced = replace_unknown(data_raw)
assert len(data_replaced) == 11, "Not correct. Keep trying."
assert isinstance(data_replaced, list), "Not correct. Keep trying."
for animal in data_replaced:
    assert isinstance(animal, dict), "Not correct. Keep trying."
    assert len(animal) == 10, "Not correct. Keep trying."

    for k, v in animal.items():
        if isinstance(v, str):
            assert "?" not in v, "Not correct. Keep trying."
        assert k in columns, "Not correct. Keep trying."
        if k == "feathers":
            assert isinstance(v, str), "Not correct. Keep trying."
        elif k == "eggs":
            assert isinstance(v, int), "Not correct. Keep trying."
        elif k == "aquatic":
            assert isinstance(v, str), "Not correct. Keep trying."

assert data_replaced[0]["feathers"] == "1", "Not correct. Keep trying."
assert data_replaced[1]["feathers"] == "0", "Not correct. Keep trying."
assert data_replaced[1]["legs"] == 4, "Not correct. Keep trying."
assert data_replaced[0]["feathers"] == "1", "Not correct. Keep trying."
assert data_replaced[8]["aquatic"] == "0", "Not correct. Keep trying."
assert data_replaced[5]["aquatic"] == "1", "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 2

The variable `data` is a list of dictionaries.

Create a function `convert_int` that:
- Receives as argument a list of dictionaries;
- In each dictionary, find all the numeric values and make sure their data type is `int`;
- Returns the modified list of dictionaires;

Hint: `isdigit()`

<!-- grade_id: cell-5b768d0072b87818, points: 0, grade: false, solution: true -->
```python
def convert_int(data):
    convert = [{w: int(w2) if isinstance(w2, str) and w2.isdigit() else w2
                for w, w2 in animal.items()}
              for animal in data]
    return convert
```

<!-- grade_id: cell-9138a4e7a89f3adc, points: 3, grade: true, solution: false -->
```python
data_replaced = data.data_replaced
columns = list(data_replaced[0].keys())
data_int = convert_int(data_replaced)

assert len(data_int) == 11, "Not correct. Keep trying."
assert isinstance(data_int, list), "Not correct. Keep trying."
for animal in data_int:
    assert isinstance(animal, dict), "Not correct. Keep trying."
    assert len(animal) == 10, "Not correct. Keep trying."

    for k, v in animal.items():
        if isinstance(v, str):
            assert "?" not in v, "Not correct. Keep trying."
        assert k in columns, "Not correct. Keep trying."
        if k in ["animal_name", "class_type"]:
            assert isinstance(v, str), "Not correct. Keep trying."
        else:
            assert isinstance(v, int), "Not correct. Keep trying."

assert data_int[0]["feathers"] == 1, "Not correct. Keep trying."
assert data_int[1]["feathers"] == 0, "Not correct. Keep trying."
assert data_int[1]["legs"] == 4, "Not correct. Keep trying."
assert data_int[0]["feathers"] == 1, "Not correct. Keep trying."
assert data_int[8]["aquatic"] == 0, "Not correct. Keep trying."
assert data_int[5]["aquatic"] == 1, "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 3

The variable `data` is a list of dictionaries.

Create a function `lower_class` that:

- Receives as argument a list of dictionaries;
- For each element of the list, find the key 'class_type' and convert the correspondent value to lower case;
- Returns the modified list of dictionaires;

<!-- grade_id: cell-0f7774d47c55eb30, points: 0, grade: false, solution: true -->
```python
def lower_class(data):
    lower = [{w: w2.lower() if w == 'class_type' else w2 
            for w, w2 in animal.items()}
            for animal in data]
    return lower
```

<!-- grade_id: cell-3a68fa6f6e905977, points: 3, grade: true, solution: false -->
```python
data_int = data.data_int
columns = list(data_int[0].keys())
data_class = lower_class(data_int)

assert len(data_class) == 11, "Not correct. Keep trying."
assert isinstance(data_class, list), "Not correct. Keep trying."
for animal in data_class:
    assert isinstance(animal, dict), "Not correct. Keep trying."
    assert len(animal) == 10, "Not correct. Keep trying."

    for k, v in animal.items():
        if isinstance(v, str):
            assert "?" not in v, "Not correct. Keep trying."
        assert k in columns, "Not correct. Keep trying."
        if k in ["animal_name", "class_type"]:
            assert isinstance(v, str), "Not correct. Keep trying."
        else:
            assert isinstance(v, int), "Not correct. Keep trying."

        if k == "class_type":
            assert v.lower() == v, "Not correct. Keep trying."

assert data_class[0]["feathers"] == 1, "Not correct. Keep trying."
assert data_class[1]["feathers"] == 0, "Not correct. Keep trying."
assert data_class[1]["legs"] == 4, "Not correct. Keep trying."
assert data_class[0]["feathers"] == 1, "Not correct. Keep trying."
assert data_class[8]["aquatic"] == 0, "Not correct. Keep trying."
assert data_class[5]["aquatic"] == 1, "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 4

The variable `data` is a list of dictionaries.

Create a function `order_animals` that:

- Receives as argument a list of dictionaries;
- Order the animals by number of legs;
- Return them as a list of (animal_name, class_type, legs) tuples;

<!-- grade_id: cell-122a9ec5cefa138a, points: 0, grade: false, solution: true -->
```python
def order_animals(data):
    data_sorted = list(sorted(utils.generate_data_lower(), key=itemgetter('legs')))
    order = [(animal['animal_name'], animal['class_type'], animal['legs'])
              for animal in data_sorted]
    return order
 
```

<!-- grade_id: cell-4c4b419744a6ebf6, points: 3, grade: true, solution: false -->
```python
data_lower = data.data_lower
data_order = order_animals(data_lower)
assert isinstance(data_order, list), "Not correct. Keep trying."
assert len(data_order) == 11, "Not correct. Keep trying."
for animal in data_order:
    assert len(animal) == 3, "Not correct. Keep trying."
    assert isinstance(animal[0], str), "Not correct. Keep trying."
    assert isinstance(animal[1], str), "Not correct. Keep trying."
    assert isinstance(animal[2], int), "Not correct. Keep trying."

assert data_order[3][0] == "pheasant", "Not correct. Keep trying."
assert data_order[6][0] == "hawk", "Not correct. Keep trying."
assert data_order[-1][0] == "giraffe", "Not correct. Keep trying."
assert data_order[0][1] == "fish", "Not correct. Keep trying."
assert data_order[4][1] == "bird", "Not correct. Keep trying."
assert data_order[-2][1] == "mammal", "Not correct. Keep trying."
assert data_order[0][2] == 0, "Not correct. Keep trying."
assert data_order[6][2] == 2, "Not correct. Keep trying."
assert data_order[-1][2] == 4, "Not correct. Keep trying."

print("Answer is correct. Good Job.")
```

Output:
```
/tmp/ipykernel_393/3764661711.py in <module>
      1 data_lower = data.data_lower
----> 2 data_order = order_animals(data_lower)
      3 assert isinstance(data_order, list), "Not correct. Keep trying."
      4 assert len(data_order) == 11, "Not correct. Keep trying."
      5 for animal in data_order:

/tmp/ipykernel_393/4051462855.py in order_animals(data)
      1 def order_animals(data):
----> 2     data_sorted = list(sorted(utils.generate_data_lower(), key=itemgetter('legs')))
      3     order = [(animal['animal_name'], animal['class_type'], animal['legs'])
      4               for animal in data_sorted]
      5     return order

NameError: name 'utils' is not defined
```

## Question 5

The variable `data` is a list of dictionaries.

Create a function `mammal_letters` that:

- Receives as argument a list of dictionaries;
- Filter all the mammals;
- Return the average number of letters in the mammals' names. The return type must be a float;

<!-- grade_id: cell-81f72032281eb943, points: 0, grade: false, solution: true -->
```python
def mammals_letters(data):
    mammals = [len(animal['animal_name'])
              for animal in data
              if animal ['class_type'] == 'mammal']
    return sum(mammals) / len(mammals)
```

<!-- grade_id: cell-03842b58f0727f29, points: 3, grade: true, solution: false -->
```python
data_lower = data.data_lower
data_letters = mammals_letters(data_lower)
assert isinstance(data_letters, float), "Not correct. Keep trying."
assert math.isclose(data_letters,
                    4.75,
                    abs_tol=0.01), "Not correct. Keep trying."
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

## Question 6

Create an `Animal` class.

The class has three class parameters: `animal_name` (str), `animal_type` (str) and `eggs` (bool).
Implement a method called `lay_eggs` that:
- Returns random number between 0 and 10 if the animal can lay eggs (eggs parameter = True)
- Returns the sentence "I can't lay eggs..." if the animal doesn't lay eggs

<!-- grade_id: cell-786ac51bf37eef4c, points: 0, grade: false, solution: true -->
```python
class Animal:
    def __init__(self, animal_name, animal_type, eggs):
        self.animal_name = animal_name
        self.animal_type = animal_type
        self.eggs = eggs
    
    def lay_eggs(self):
        return random.randint(0,10) if self.eggs else "I can't lay eggs..."
```

<!-- grade_id: cell-f50a68ec0e515188, points: 5, grade: true, solution: false -->
```python
test1 = Animal("Duck", "bird", True)
assert isinstance(test1.lay_eggs(), int), "Not correct. Keep trying."
assert test1.lay_eggs() >= 0, "Not correct. Keep trying."
assert test1.lay_eggs() <= 10, "Not correct. Keep trying."

test2 = Animal("Dolphin", "mammal", False)
assert isinstance(test2.lay_eggs(), str), "Not correct. Keep trying."
assert test2.lay_eggs() == "I can't lay eggs...", "Not correct. Keep trying."
print("Answer is correct. Good Job.")
```

Output:
```
Answer is correct. Good Job.
```

### Last but not least, get your grade!

To know your grade, run the following command in the **COMMAND LINE:**

`python grading.py`